# Functions

Topics:
1. How to declare a function
2. Duck-typing in Julia
3. Mutating vs. non-mutating functions
4. Some higher order functions

## How to declare a function
Julia gives us a few different ways to write a function. The first requires the `function` and `end` keywords

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

Output:
```
sayhi (generic function with 1 method)
```

```julia
function f(x)
    x^2
end
```

Output:
```
f (generic function with 1 method)
```

We can call either of these functions like this:

```julia
sayhi("C-3PO")
```

Output:
```
Hi C-3PO, it's great to see you!
```

```julia
f(42)
```

Output:
```
1764
```

Alternatively, we could have declared either of these functions in a single line

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

```julia
f2(x) = x^2
```

```julia
sayhi2("R2D2")
```

```julia
f2(42)
```

Finally, we could have declared these as "anonymous" functions

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

Output:
```
#3 (generic function with 1 method)
```

```julia
f3 = x -> x^2
```

Output:
```
#5 (generic function with 1 method)
```

```julia
sayhi3("Chewbacca")
```

Output:
```
Hi Chewbacca, it's great to see you!
```

```julia
f3(42)
```

Output:
```
1764
```

## Duck-typing in Julia
*"If it quacks like a duck, it's a duck."* <br><br>
Julia functions will just work on whatever inputs make sense. <br><br>
For example, `sayhi` works on the name of this minor tv character, written as an integer...

```julia
sayhi(55595472)
```

Output:
```
Hi 55595472, it's great to see you!
```

And `f` will work on a matrix. 

```julia
A = rand(3, 3)
A
```

Output:
```
3×3 Array{Float64,2}:
 0.605761   0.461066   0.605301
 0.381813   0.0882238  0.66921
 0.0688605  0.0380225  0.942272
```

```julia
f(A)
```

Output:
```
3×3 Array{Float64,2}:
 0.584669  0.342988   1.24558
 0.311054  0.209269   0.92073
 0.121116  0.0709313  0.955004
```

`f` will also work on a string like "hi" because `*` is defined for string inputs as string concatenation.

```julia
f("hi")
```

Output:
```
"hihi"
```

On the other hand, `f` will not work on a vector. Unlike `A^2`, which is well-defined, the meaning of `v^2` for a vector, `v`, is not a well-defined algebraic operation. 

```julia
v = rand(3)
```

Output:
```
3-element Array{Float64,1}:
 0.07095868735570066
 0.32896765407585615
 0.9001343925919993
```

```julia
f(v)
```

Output:
```
[3] f(::Array{Float64,1}) at .\In[2]:2
 [4] top-level scope at In[16]:1
 [5] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

## Mutating vs. non-mutating functions

By convention, functions followed by `!` alter their contents and functions lacking `!` do not.

For example, let's look at the difference between `sort` and `sort!`.


```julia
v = [3, 5, 2]
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

```julia
sort(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

`sort(v)` returns a sorted array that contains the same elements as `v`, but `v` is left unchanged. <br><br>

On the other hand, when we run `sort!(v)`, the contents of v are sorted within the array `v`.

```julia
sort!(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

## Some higher order functions

### map

`map` is a "higher-order" function in Julia that *takes a function* as one of its input arguments. 
`map` then applies that function to every element of the data structure you pass it. For example, executing

```julia
map(f, [1, 2, 3])
```
will give you an output array where the function `f` has been applied to all elements of `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

Here we've squared all the elements of the vector `[1, 2, 3]`, rather than squaring the vector `[1, 2, 3]`.

To do this, we could have passed to `map` an anonymous function rather than a named function, such as

```julia
x -> x^3
```

Output:
```
#7 (generic function with 1 method)
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  1
  8
 27
```

and now we've cubed all the elements of `[1, 2, 3]`!

### broadcast

`broadcast` is another higher-order function like `map`. `broadcast` is a generalization of `map`, so it can do every thing `map` can do and more. The syntax for calling `broadcast` is the same as for calling `map`

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

and again, we've applied `f` (squared) to all the elements of `[1, 2, 3]` - this time by "broadcasting" `f`!

Some syntactic sugar for calling `broadcast` is to place a `.` between the name of the function you want to `broadcast` and its input arguments. For example,

```julia
broadcast(f, [1, 2, 3])
```
is the same as
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
 1
 4
 9
```

```julia
f([1, 2, 3])
```

Output:
```
[3] f(::Array{Int64,1}) at .\In[2]:2
 [4] top-level scope at In[34]:1
 [5] include_string(::Function, ::Module, ::String, ::String) at .\loading.jl:1091
```

Notice again how different this is from calling 
```julia
f([1, 2, 3])
```
We can square every element of a vector, but we can't square a vector!

To drive home the point, let's look at the difference between

```julia
f(A)
```
and
```julia
f.(A)
```
for a matrix `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Array{Int64,2}:
  30   36   42
  66   81   96
 102  126  150
```

As before we see that for a matrix, `A`,
```
f(A) = A^2 = A * A
``` 

On the other hand,

```julia
B = f.(A)
```

Output:
```
3×3 Array{Int64,2}:
  1   4   9
 16  25  36
 49  64  81
```

contains the squares of all the entries of `A`.

This dot syntax for broadcasting allows us to write relatively complex compound elementwise expressions in a way that looks natural/closer to mathematical notation. For example, we can write

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Array{Float64,2}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

instead of

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

Output:
```
3×3 Array{Float64,2}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

and the two will perform exactly the same.

### Exercises

#### 6.1 
Write a function `add_one` that adds 1 to its input.

```julia
add_one = input -> input+= 1
```

Output:
```
#18 (generic function with 1 method)
```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Use `map` or `broadcast` to increment every element of matrix `A` by `1` and assign it to a variable `A1`.

```julia
A1 = A .+ 1
```

Output:
```
3×3 Array{Int64,2}:
 2  3   4
 5  6   7
 8  9  10
```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 
Use the broadcast dot syntax to increment every element of matrix `A1` by `1` and store it in variable `A2`

```julia
A2 = A1 .+ 1
```

Output:
```
3×3 Array{Int64,2}:
 3   4   5
 6   7   8
 9  10  11
```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

Please click on `Validate` on the top, once you are done with the exercises.